# 找到奇偶异常值

> 题目来自 CodeWars [Find The Parity Outlier](https://www.codewars.com/kata/find-the-parity-outlier)

给定一个长度至少为 3 整数列表，其中要么所有数都是奇数，只有一个数是偶数（即本题所说的奇偶异常值）；要么所有数都是偶数，唯独有一个奇数异常值。

完成函数 `find_outlier`，返回列表中的这个“奇偶异常值”。

输入 / 输出示例见下方的 `assert` cell。

请从 `# YOUR CODE HERE` 下方另起一行填写你的答案。

<!-- grade_id: finder, points: 0, grade: false, solution: true -->
```python
def find_outlier(lst):
    # YOUR CODE HERE
    raise NotImplementedError()
```

提交前请你确保下列 cell 运行时不报错。

<!-- grade_id: can_find_odd, points: 50, grade: true, solution: false -->
```python
assert find_outlier([2, 4, 6, 8, 10, 3]) == 3
assert find_outlier([2, 4, 0, 100, 4, 11, 2602, 36]) == 11
```

<!-- grade_id: can_find_even, points: 50, grade: true, solution: false -->
```python
assert find_outlier([1, 3, 5, 7, 9, 2]) == 2
assert find_outlier([160, 3, 1719, 19, 11, 13, -21]) == 160
```