# 嵌套列表排序

> 题目来自 CodeWars [Sort only integers in Nested List](https://www.codewars.com/kata/sort-only-integers-in-nested-list/train/python)

完成函数 `sort_nested_list(nlst)`，参数 `nlst` 是一个嵌套列表，你需要保持 `nlst` 中的嵌套结构不变，只给列表中的整数从小到大排序。例如：

```Python
Input   : [ [ [ 2, 1 ], [ 4, 3 ] ], [ [ 6, 5 ], [ 8, 7 ] ] ]
Output  : [ [ [ 1, 2 ], [ 3, 4 ] ], [ [ 5, 6 ], [ 7, 8 ] ] ]
    
Input   : [ [ [ 82, 87 ], [ 47, 44 ] ], [ [ 70, 94 ], [ 40, 64 ] ] ]
Output  :[ [ [ 40, 44 ], [ 47, 64 ] ], [ [ 70, 82 ], [ 87, 94 ] ] ]
```

返回排序后的嵌套列表。

请从 `# YOUR CODE HERE` 下方另起一行填写你的答案。

**注意**，不要删除 `# YOUR CODE HERE` 注释。

<!-- grade_id: sort, points: 0, grade: false, solution: true -->
```python
def sort_nested_list(nlst):
    ### BEGIN SOLUTION
    def fl(l):
        for x in l:
            if isinstance(x,list):
                for j in fl(x):
                    yield j
            else:
                yield x

    numbers = iter(sorted(fl(nlst)))
    def b(n, nlst):
        return [next(n) if isinstance(c, int)else b(n, c) for c in nlst]
    return b(numbers, nlst)
    ### END SOLUTION
```

提交前请你确保下列 cell 运行时不报错。

<!-- grade_id: simple_nest, points: 50, grade: true, solution: false -->
```python
A=[[[3, 2, 1]]]
expected=[[[1, 2, 3]]]
assert sort_nested_list(A) == expected
### BEGIN HIDDEN TESTS
A=[[[3, 2, 1]]]
expected=[[[1, 2, 3]]]
assert sort_nested_list(A) == expected

A=[[[15, 30, 92]]]
E2=[[[15, 30, 92]]]
assert sort_nested_list(A) == E2
### END HIDDEN TESTS
```

<!-- grade_id: complex_nest, points: 50, grade: true, solution: false -->
```python
A = [ [ [ 82, 87 ], [ 47, 44 ] ], [ [ 70, 94 ], [ 40, 64 ] ] ]
expected = [ [ [ 40, 44 ], [ 47, 64 ] ], [ [ 70, 82 ], [ 87, 94 ] ] ]
assert sort_nested_list(A) == expected
### BEGIN HIDDEN TESTS
A=[[[82, 87], [47, 44]], [[70, 94], [40, 64]]]
expected=[[[40, 44], [47, 64]], [[70, 82], [87, 94]]]
assert sort_nested_list(A) == expected

A=[[[29, 32], [82, 61], [75, 91]], [[69, 99], [74, 23], [70, 97]]]
E4=[[[23, 29], [32, 61], [69, 70]], [[74, 75], [82, 91], [97, 99]]]
assert sort_nested_list(A) == E4

A=[[[]]]
expected=[[[]]]
assert sort_nested_list(A) == expected
### END HIDDEN TESTS
```