# 数根

> 题目来自 CodeWars [Sum of Digits / Digital Root](https://www.codewars.com/kata/sum-of-digits-slash-digital-root)
>
> 数根的定义来自 [数根 - 维基百科](https://zh.wikipedia.org/w/index.php?title=數根&oldid=42879432)

完成函数 `digital_root`，给定一个正整数，返回它的数根。

数根是将一正整数的各个位数相加（即横向相加），若加完后的值大于10的话，则继续将各位数进行横向相加直到其值小于十为止。如：

```
digital_root(16)
=> 1 + 6
=> 7

digital_root(942)
=> 9 + 4 + 2
=> 15 ...
=> 1 + 5
=> 6

digital_root(132189)
=> 1 + 3 + 2 + 1 + 8 + 9
=> 24 ...
=> 2 + 4
=> 6

digital_root(493193)
=> 4 + 9 + 3 + 1 + 9 + 3
=> 29 ...
=> 2 + 9
=> 11 ...
=> 1 + 1
=> 2
```

请从 `# YOUR CODE HERE` 下方另起一行填写你的答案。

<!-- grade_id: root, points: 0, grade: false, solution: true -->
```python
def digital_root(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

提交前请你确保下列 cell 运行时不报错。

<!-- grade_id: calc_once, points: 50, grade: true, solution: false -->
```python
assert digital_root(16) == 7
```

<!-- grade_id: calc_multi_times, points: 50, grade: true, solution: false -->
```python
assert digital_root(942) == 6
assert digital_root(132189) == 6
assert digital_root(493193) == 2
```