# 不要再转我的剑了！

> 题目来自 CodeWars [Stop gninnipS My sdroW!](https://www.codewars.com/kata/stop-gninnips-my-sdrow)

完成函数 `spin_words`，给定一句只包含字母和空格的英文句子，将长度为 5 或以上的单词倒转过来，如：

```
spin_words("Hey fellow warriors" ) => "Hey wollef sroirraw" 
spin_words("This is a test") => "This is a test" 
spin_words("This is another test" ) => "This is rehtona test"
```

请从 `# YOUR CODE HERE` 下方另起一行填写你的答案。

**注意**，不要删除 `# YOUR CODE HERE` 注释。

<!-- grade_id: spin, points: 0, grade: false, solution: true -->
```python
def spin_words(sentence):
    # YOUR CODE HERE
    raise NotImplementedError()
```

提交前请你确保下列 cell 运行时不报错。

<!-- grade_id: spin_one_word, points: 50, grade: true, solution: false -->
```python
assert spin_words('Welcome') == 'emocleW'
```

<!-- grade_id: spin_multi_words, points: 50, grade: true, solution: false -->
```python
assert spin_words( "Hey fellow warriors" ) == "Hey wollef sroirraw" 
assert spin_words( "This is a test") == "This is a test" 
assert spin_words( "This is another test" ) == "This is rehtona test"
```

```python

```