# 按字母顺序排列

> 题目来自 CodeWars [Where is my parent!?(cry)](https://www.codewars.com/kata/where-is-my-parent-cry/solutions/python/me/best_practice)

完成函数 `ab_sort(s)`，参数 `s` 是只包含大写与小写英文字母的字符串，返回排序后的 `s`。

排序方法：首先按拉丁字母表顺序从小到大排序，然后对于相同的字母，大写在前，小写在后排序，如：`'bBBAaaCCc' -> 'AaaBBbCCc'`

请从 `# YOUR CODE HERE` 下方另起一行填写你的答案。

**注意**，不要删除 `# YOUR CODE HERE` 注释。

<!-- grade_id: ab_sort, points: 0, grade: false, solution: true -->
```python
def ab_sort(s):
    ### BEGIN SOLUTION
    return ''.join(sorted(s, key=lambda c: (c.lower(), c)))
    ### END SOLUTION
```

提交前请你确保下列 cell 运行时不报错。

<!-- grade_id: test, points: 100, grade: true, solution: false -->
```python
assert ab_sort("abBA") == "AaBb"
### BEGIN HIDDEN TESTS
assert ab_sort("AaaaaZazzz") == "AaaaaaZzzz"
assert ab_sort("CbcBcbaA") == "AaBbbCcc"
assert ab_sort("xXfuUuuF") == "FfUuuuXx"
assert ab_sort("") == ""
### END HIDDEN TESTS
```

```python

```