# 如果超过 n 次，删！

> 题目来自 CodeWars [Delete occurrences of an element if it occurs more than n times](https://www.codewars.com/kata/delete-occurrences-of-an-element-if-it-occurs-more-than-n-times)

完成函数 `delete_nth(lst, n)`，参数 `lst` 是一个整数列表，确保任何数字出现次数不超过 n 次。

如果列表中有一个数出现超过 n 次，将它超过的部分删除。比如：


```Python
delete_nth ([1, 1, 1, 1], 2)  # return [1, 1]
delete_nth ([20, 37, 20, 21], 1)  # return [20, 37, 21]
```

请从 `# YOUR CODE HERE` 下方另起一行填写你的答案。

**注意**，不要删除 `# YOUR CODE HERE` 注释。

<!-- grade_id: delete, points: 0, grade: false, solution: true -->
```python
def delete_nth(lst, n):
    ### BEGIN SOLUTION
    ans = []
    for o in lst:
        if ans.count(o) < n: ans.append(o)
    return ans
    ### END SOLUTION
```

提交前请你确保下列 cell 运行时不报错。

<!-- grade_id: only_ones, points: 20, grade: true, solution: false -->
```python
assert delete_nth([1, 1, 1, 1], 2) == [1, 1]
### BEGIN HIDDEN TESTS
assert delete_nth([1, 1, 1, 1], 1) == [1]
assert delete_nth([1, 1, 1, 1], 2) == [1, 1]
assert delete_nth([1, 1, 1, 1], 3) == [1, 1, 1]
### END HIDDEN TESTS
```

<!-- grade_id: ones_and_others, points: 30, grade: true, solution: false -->
```python
assert delete_nth ([20, 37, 20, 21], 1) == [20,37,21]
### BEGIN HIDDEN TESTS
assert delete_nth ([20, 37, 20, 21], 1) == [20,37,21]
assert delete_nth ([20, 37, 20, 21], 2) == [20,37,20,21]
### END HIDDEN TESTS
```

<!-- grade_id: handle_complex, points: 50, grade: true, solution: false -->
```python
assert delete_nth ([1, 1, 3, 3, 7, 2, 2, 2, 2], 3) == [1, 1, 3, 3, 7, 2, 2, 2]
### BEGIN HIDDEN TESTS
assert delete_nth ([1,1,3,3,7,2,2,2,2], 4) == [1, 1, 3, 3, 7, 2, 2, 2, 2]
assert delete_nth ([1,1,3,3,7,2,2,2,2], 3) == [1, 1, 3, 3, 7, 2, 2, 2]
assert delete_nth ([1,1,3,3,7,2,2,2,2], 2) == [1, 1, 3, 3, 7, 2, 2]
assert delete_nth ([1,1,3,3,7,2,2,2,2], 1) == [1, 3, 7, 2]
### END HIDDEN TESTS
```

```python

```