# “查重”编码器

> 题目来自 CodeWars [Duplicate Encoder](https://www.codewars.com/kata/duplicate-encoder)

完成函数 `duplicate_encode(word)`，对字符串 `word` 进行编码，对于字符串中的每一个字符：

* 如果字符出现次数为 1，把字符编码为 `(`
* 如果字符出现不止 1 次，把字符编码为 `)`

返回编码后的字符串。同一字母的大小写视为相同的字符，如 `Aa` 视为 `a` 出现 2 次，编码为 `))`。

请从 `# YOUR CODE HERE` 下方另起一行填写你的答案。

**注意**，不要删除 `# YOUR CODE HERE` 注释。

<!-- grade_id: encode, points: 0, grade: false, solution: true -->
```python
def duplicate_encode(word):
    ### BEGIN SOLUTION
    lower_word = word.lower()
    return ''.join('(' if lower_word.count(char) == 1 else ')' for char in lower_word)
    ### END SOLUTION
```

提交前请你确保下列 cell 运行时不报错。

<!-- grade_id: lowercase, points: 20, grade: true, solution: false -->
```python
assert duplicate_encode("din") == "((("
assert duplicate_encode("dad") == ")()"
### BEGIN HIDDEN TESTS
assert duplicate_encode("din") == "((("
assert duplicate_encode("big") == "((("
assert duplicate_encode("dad") == ")()"
assert duplicate_encode("bab") == ")()"
### END HIDDEN TESTS
```

<!-- grade_id: mixcase, points: 30, grade: true, solution: false -->
```python
assert duplicate_encode("Recede") == "()()()"
### BEGIN HIDDEN TESTS
assert duplicate_encode("Recede") == "()()()"
assert duplicate_encode("Big") == "((("
assert duplicate_encode("Dad") == ")()"
### END HIDDEN TESTS
```

<!-- grade_id: complex, points: 50, grade: true, solution: false -->
```python
assert duplicate_encode("Success") == ")())())"
assert duplicate_encode("(( @") == "))(("
### BEGIN HIDDEN TESTS
assert duplicate_encode("Successs") == ")())()))"
assert duplicate_encode("(( @") == "))(("
### END HIDDEN TESTS
```

```python

```